Прежде чем сдать эту работу, убедитесь, что всё выполняется так, как ожидается.
Сначала перезапустите ядро (в меню выберите Kernel → Restart), а затем запустите все ячейки (в меню выберите Cell → Run All).

Убедитесь, что вы заполнили все места, где указано YOUR CODE HERE или "YOUR ANSWER HERE", а также указали своё имя ниже.

```python
NAME = ""
```

Убедитесь, что используете версию Python > 3. Все тетрадки отрабатывают точно на версии 3.10.10.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

В этой тетрадке вам предлагается реализовать задания при помощи библиотеки numpy. Крайне желательно при выполнении разных заданий не использовать циклы. Если вы не понимаете, как выполнить конкретную задачу без циклов - убирайте их постепенно, разбейте выполнение задачи на этапы и постепенно векторизуйте каждый этап.

Если вам непонятно, как выполнить ту или иную операцию - не стесняйтесь гуглить, если над вами повис дедлайн, а вы всё ещё на чём-то застряли - пишите преподу/в чат, это ок.

**Обсудить решение лучше и сделать в итоге хоть что-то гораздо лучше, чем скатать чужой код или сдать пустую тетрадку.**

```python
import numpy as np
```

```python
np.__version__
```

#### Задание 1. (0.4 балла)

Инициализируйте нулями массивы следующих форм и верните их в функции:

* (5)
* (5, 1)
* (10, 2)
* (3, 6, 2)
* (4, 3, 2, 1)

<!-- grade_id: cell-7b463169be3cf4d3, points: 0, grade: false, solution: true -->
```python
def init_zeros() -> tuple:
    # YOUR CODE HERE
    raise NotImplementedError()
    return a, b, c, d, e
```

<!-- grade_id: cell-ed2b2b28c486eddd, points: 1, grade: true, solution: false -->
```python
arrays = init_zeros()
assert len(arrays) == 5
```

#### Задание 2. (0.4 балла)

Инициализируйте единицами массивы таких же форм и приведите их к целочисленному виду:

<!-- grade_id: cell-d9e2df1bb85b7a88, points: 0, grade: false, solution: true -->
```python
def init_ones() -> tuple:
    # YOUR CODE HERE
    raise NotImplementedError()
    return a, b, c, d, e
```

<!-- grade_id: cell-bbd2a3e37a90f667, points: 1, grade: true, solution: false -->
```python
arrays = init_ones()
assert len(arrays) == 5
```

#### Задание 3. (0.4 балла)

Выполните для векторов выше `reshape` в "форму"
* с двумя компонентами
* с тремя компонентами
* с четырьмя компонентами


<!-- grade_id: cell-e10a49d33718b6dc, points: 0, grade: false, solution: true -->
```python
def reshape_tensor_first(arrays: tuple) -> tuple:
    a = arrays[0] # (5,)
    a_two_components = a.reshape(5, 1)
    a_three_components = a.reshape(5, 1, 1)
    a_four_components = a.reshape(5, 1, 1, 1)
    return a_two_components, a_three_components, a_four_components

def reshape_tensor_second(arrays: tuple) -> tuple:
    b = arrays[1] # (5, 1)
    # YOUR CODE HERE
    raise NotImplementedError()
    return b_two_components, b_three_components, b_four_components

def reshape_tensor_third(arrays: tuple) -> tuple:
    c = arrays[2] # (10, 2)
    # YOUR CODE HERE
    raise NotImplementedError()
    return c_two_components, c_three_components, c_four_components

def reshape_tensor_fourth(arrays: tuple) -> tuple:
    d = arrays[3] # (3, 6, 2)
    # YOUR CODE HERE
    raise NotImplementedError()
    return d_two_components, d_three_components, d_four_components

def reshape_tensor_fifth(arrays: tuple) -> tuple:
    e = arrays[4] # (4, 3, 2, 1)
    # YOUR CODE HERE
    raise NotImplementedError()
    return e_two_components, e_three_components, e_four_components
```

<!-- grade_id: cell-f7e54eeb544c26bd, points: 1, grade: true, solution: false -->
```python
a_arrays = reshape_tensor_first(arrays)
b_arrays = reshape_tensor_second(arrays)
c_arrays = reshape_tensor_third(arrays)
d_arrays = reshape_tensor_fourth(arrays)
e_arrays = reshape_tensor_fifth(arrays)

assert len(a_arrays) == 3
for i in range(len(a_arrays)):
    assert np.prod(a_arrays[i].shape) == 5

```

#### Задание 4 (0.4 балла)
Напишите функцию get_info, которая возвращает в сторогом порядке:
1) максимум по 1 оси
2) минимум по 0 оси
3) сумму по всем осям
4) среднее по 0 оси
5) медиану по 1 оси

<!-- grade_id: cell-ebd5f6b3831263b4, points: 0, grade: false, solution: true -->
```python
def get_info(v: np.ndarray) -> tuple:
    # YOUR CODE HERE
    raise NotImplementedError()
    return max_1, min_0, sum_all, mean_0, median_1
```

<!-- grade_id: cell-85061c8377a463dc, points: 1, grade: true, solution: false -->
```python
v = np.random.randint(0, 9, size=(5, 7))
result = get_info(v)
assert len(result) == 5
assert result[0].shape == (5,)
assert result[1].shape == (7,)
assert result[2].shape == ()
assert result[3].shape == (7,)
assert result[4].shape == (5,)


```

#### Задание 5 (0.4 балла)
Напишите функции, которые принимают на вход вектор и возвращают:
1) все числа больше 5
2) все числа меньше 7 в обратном порядке
3) все числа больше 2, которые делятся на 4 без остатка
4) все числа, остаток от деления на 5 которых равен 3
5) индекс максимума

<!-- grade_id: cell-243a77fb2657b7c9, points: 0, grade: false, solution: true -->
```python
def get_more_than_5(v: np.ndarray) -> np.ndarray:
    # YOUR CODE HERE
    raise NotImplementedError()

def get_less_than_7_in_reverse_order(v: np.ndarray) -> np.ndarray:
    # YOUR CODE HERE
    raise NotImplementedError()

def get_more_than_2_divisible_by_4(v: np.ndarray) -> np.ndarray:
    # YOUR CODE HERE
    raise NotImplementedError()

def get_remainder_3(v: np.ndarray) -> np.ndarray:
    # YOUR CODE HERE
    raise NotImplementedError()

def get_max_index(v: np.ndarray) -> int:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e377b625472203d3, points: 1, grade: true, solution: false -->
```python
v = np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]])

assert get_more_than_5(v).shape == (4,)
assert get_less_than_7_in_reverse_order(v).shape == (6,)
assert get_more_than_2_divisible_by_4(v).shape == (2,)
assert get_remainder_3(v).shape == (2,)
assert get_max_index(v) == 8

```

#### Задание 6 (0.4 балла)
Напишите функцию make_magic, которая выполняет следующие операции над вектором последовательно:
1) умножьте его на 3
2) разделите на 7
3) возведите в степень $\frac{2}{3}$
4) посчитайте $\sin$ от него
5) посчитайте $\log$ от него

<!-- grade_id: cell-40398a00bda1bde6, points: 0, grade: false, solution: true -->
```python
def make_magic(v: np.ndarray) -> np.ndarray:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-73081c2cfa185625, points: 1, grade: true, solution: false -->
```python
v = np.random.randint(0, 10, size=(5, 7))
assert make_magic(v).shape == v.shape


```

#### Задание 7 (0.4 балла)
Напишите функцию, которая возвращает массив размера 3 x 3 из нечетных чисел

<!-- grade_id: cell-ba61dfa73a408285, points: 0, grade: false, solution: true -->
```python
def get_odd_numbers() -> np.ndarray:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a458843ad032c481, points: 1, grade: true, solution: false -->
```python
v = get_odd_numbers()
assert v.shape == (3, 3)
assert np.all(v % 2 == 1)
```

#### Задание 8 (0.4 балла)
Напишите функцию, которая возвращает массив размера 4 х 4 из четных чисел в обратном порядке

<!-- grade_id: cell-4aeab8992dca985f, points: 0, grade: false, solution: true -->
```python
def get_even_numbers() -> np.ndarray:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-093abd3c29aac7ed, points: 1, grade: true, solution: false -->
```python
v = get_even_numbers()
assert v.shape == (4, 4)
assert np.all(v % 2 == 0)


```

#### Задание 9 (1,6 баллов)
Напишите функцию, которая возвращает координатную сетку размера 11 х 11 с точкой (0, 0) в центре и шагом 0.1 между клетками

<!-- grade_id: cell-4856f01275ff348c, points: 0, grade: false, solution: true -->
```python
def get_grid() -> np.ndarray:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2bb420a5d4806201, points: 4, grade: true, solution: false -->
```python
grid = get_grid()

assert grid.shape == (11, 11, 2)
assert np.all(grid[0][0] == np.array([-0.5, 0.5]))
assert np.all(grid[5][5] == np.array([0, 0]))
assert np.all(grid[10][10] == np.array([0.5, -0.5]))

```

#### Задание 10 (0.8 баллов)
Напишите функцию get_chess, которая возвращает сетку размера 8 х 8, в которой 0 и 1 стоит в шахматном порядке. Клетка с индексом (0, 0) равна 1.

<!-- grade_id: cell-a801d62beae0b0ad, points: 0, grade: false, solution: true -->
```python
def get_chess() -> np.ndarray:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1fb1fcd1fef7bbd9, points: 2, grade: true, solution: false -->
```python
chess = get_chess()

assert chess.shape == (8, 8)
assert np.all(chess[0, 0] == 1)

```

#### Задание 11 (0.4 балла)
Напишите функцию, которая возвращает нижне-треугольную матрицу переменного размера (передается в функцию). То есть это единичная матрица, замощенная 1 под главной диагональю.

<!-- grade_id: cell-3e3388a06d9a4862, points: 0, grade: false, solution: true -->
```python
def get_lower_triangle(n: int, m: int) -> np.ndarray:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2ec1eb3d867b9591, points: 1, grade: true, solution: false -->
```python
a = get_lower_triangle(3, 3)

assert a.shape == (3, 3)
assert a[1][1] == 1

```

#### Задание 12 (0.4 балла)
Сделать то же самое, но теперь единицы над главной диагональю

<!-- grade_id: cell-a2d6f18e436976f2, points: 0, grade: false, solution: true -->
```python
def get_upper_triangle(n: int, m: int) -> np.ndarray:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-42d65868aae6e7be, points: 1, grade: true, solution: false -->
```python
a = get_upper_triangle(3, 3)

assert a.shape == (3, 3)
assert a[1][1] == 1


```

#### Задание 13 (0.8 балла)
Дана выборка $\{X_i\}$ некоторой случайной величины $X$, необходимо оценить
* Математическое ожидание $E[X] = \frac{1}{N}\sum_{i=1}^N X_i$ этой случайной величины
* Дисперсию $D[X] = \frac{1}{N}\sum_{i=1}^N (X_i - E[X])^2$ этой случайной величины

<!-- grade_id: cell-ef516aa8bc309059, points: 0, grade: false, solution: true -->
```python
def expected_value(X: np.ndarray) -> float:
    # YOUR CODE HERE
    raise NotImplementedError()

def dispersion(X: np.ndarray) -> float:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-473a046662c90717, points: 2, grade: true, solution: false -->
```python
a = np.arange(10)
expect = expected_value(a)
disp = dispersion(a)

assert expect == 4.5
assert disp == 8.25

```

#### Задание 14 (0.8 баллов)
Напишите функцию, которая принимает на вход два вектора и считает среднее арифметическое квадратов разностей их значений.

<!-- grade_id: cell-cfdff612cb1b0e4f, points: 0, grade: false, solution: true -->
```python
def mean_square_difference(y1: np.ndarray, y2: np.ndarray) -> float:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6e36f77a08350606, points: 2, grade: true, solution: false -->
```python
a = np.arange(10)
b = np.arange(10)
assert mean_square_difference(a, b) == 0


```

#### Задание 15 (2 балла)
Даны два одинаковой длины вектора (списка) y₁ и y₂, каждый из которых состоит только из 0 и 1.
Мы хотим найти такие границы разделения (индексы), чтобы если "разрезать" эти векторы в этом месте, выполнялось следующее условие:

Левая часть (все элементы до выбранного индекса) — содержит больше совпадающих нулей, чем единиц.
То есть, среди элементов, где y₁[i] == y₂[i], нулей больше, чем единиц.

Правая часть (все элементы после выбранного индекса) — наоборот, содержит больше совпадающих единиц, чем нулей.

Если какая-то из частей оказывается пустой (например, весь массив слева или справа), то она автоматически считается подходящей.

```python
def find_thresholds(y1, y2, debug=False):
    assert len(y1) == len(y2) and len(y1) > 0
    n = len(y1)
    matching_zeros = [1 if y1[i] == y2[i] == 0 else 0 for i in range(n)]
    matching_ones = [1 if y1[i] == y2[i] == 1 else 0 for i in range(n)]
    res = []
    if sum(matching_ones) / n > 0.5:
        res.append(0)
    for i in range(1, n):
        if debug: # если вам непонятна логика работы этой функции, передайте третьим аргументом True
            print(matching_zeros[:i], matching_ones[i:], i)
        if sum(matching_zeros[:i]) / i > 0.5 and sum(matching_ones[i:]) / (n - i) > 0.5:
            res.append(i)
    if sum(matching_zeros) / n > 0.5:
        res.append(n)
    return res
```

<!-- grade_id: cell-0d406185c33b96e0, points: 0, grade: false, solution: true -->
```python
def find_thresholds_np(y1: np.ndarray, y2: np.ndarray) -> np.ndarray:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4cae24ac94977bba, points: 5, grade: true, solution: false -->
```python
np.testing.assert_array_equal(find_thresholds([0], [0]), find_thresholds_np(np.array([0]), np.array([0])))
np.testing.assert_array_equal(find_thresholds([1], [1]), find_thresholds_np(np.array([1]), np.array([1])))
np.testing.assert_array_equal(
    find_thresholds([0, 0, 1, 1], [0, 0, 1, 1]), 
    find_thresholds_np(np.array([0, 0, 1, 1]), np.array([0, 0, 1, 1]))
)
np.testing.assert_array_equal(
    find_thresholds([0, 0, 1, 0, 0, 1, 1, 1], [0, 0, 0, 0, 0, 1, 1, 1]),
    find_thresholds_np(np.array([0, 0, 1, 0, 0, 1, 1, 1]), np.array([0, 0, 0, 0, 0, 1, 1, 1]))
)
```

## Фидбек

Оставьте здесь свои впечатления от домашки: 
- что понравилось больше всего?
- что можно улучшить?
- как часто прибегали к документации \ помощи ChatGPT \ однокурсников.
